#include "singly_linked_list.h"

#include <stdio.h>
#include <stdlib.h>

// 全局头指针
static link head = NULL;

// 创建新节点
link make_node(unsigned char item) {
    // TODO: 在这里添加你的代码
    // I AM NOT DONE
}

// 释放节点
void free_node(link p) { free(p); }

// 查找节点
link search(unsigned char key) {
    // TODO: 在这里添加你的代码
    // I AM NOT DONE
}

// 在链表头部插入节点
void insert(link p) {
    // TODO: 在这里添加你的代码
    // I AM NOT DONE
}

// 删除指定节点
void delete(link p) {
    // TODO: 在这里添加你的代码
    // I AM NOT DONE
}

// 遍历链表
void traverse(void (*visit)(link)) {
    // TODO: 在这里添加你的代码
    // I AM NOT DONE
}

// 销毁整个链表
void destroy(void) {
    // TODO: 在这里添加你的代码
    // I AM NOT DONE
}

// 在链表头部推入节点
void push(link p) { 
    // TODO: 在这里添加你的代码
    // I AM NOT DONE
 }

// 从链表头部弹出节点
link pop(void) {
    // TODO: 在这里添加你的代码
    // I AM NOT DONE
}

// 释放链表内存
void free_list(link list_head) {
    // TODO: 在这里添加你的代码
    // I AM NOT DONE
}
